# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

```python
import pandas as pd
import numpy as np
df = pd.read_csv("assets/NISPUF17.csv", index_col = 0 )

```

```python
# df["EDUC1"]
len(df)
```

Output:
```
28465
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    less = (len(df[df["EDUC1"]==1])) / len(df)
    high_school = (len(df[df["EDUC1"]==2])) / len(df)
    not_college =  (len(df[df["EDUC1"]==3])) / len(df)
    college =  (len(df[df["EDUC1"]==4])) / len(df)
    dict = {"less than high school":less,
    "high school":high_school,
    "more than high school but not college":not_college,
    "college":college}
    return dict
    raise NotImplementedError()
```

```python
type(proportion_of_education())
```

Output:
```
dict
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
df["SEX"].values
```

Output:
```
array([1, 1, 2, ..., 2, 2, 2])
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    df_breast_feeded = df[df["CBF_01"]==1]
    df_not_breast_feeded = df[df["CBF_01"]==2]
    breast_feeded_mean = df_breast_feeded.loc[:, "P_NUMFLU"].mean()
    not_breast_feeded_mean = df_not_breast_feeded.loc[:, "P_NUMFLU"].mean()
    tuple = (breast_feeded_mean, not_breast_feeded_mean)
    return tuple
    raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
import re
cols_list = df.columns.tolist()
cols_str = ' '.join(i for i in cols_list)
print(len(re.findall("VRC", cols_str)))
find_vrc = re.findall("\s[0-9a-zA-Z_]*VRC[0-9a-zA-Z_]*\s", cols_str)
print(re.findall("\s[0-9a-zA-Z_]*VRC[0-9a-zA-Z_]*\s", cols_str))
print(len(re.findall("\s[0-9a-zA-Z_]*VRC[0-9a-zA-Z_]*\s", cols_str)))
vrc_list = [str[1:-1] for str in find_vrc]
print(vrc_list)
df_vrc = df[vrc_list]
print(df_vrc.head())
```

Output:
```
29
[' P_NUMVRC ', ' P_U12VRC ', ' DVRC1 ', ' DVRC3 ', ' DVRC5 ', ' DVRC7 ', ' DVRC9 ', ' VRC1_AGE ', ' VRC3_AGE ', ' VRC5_AGE ', ' VRC7_AGE ', ' VRC9_AGE ', ' XVRCTY1 ', ' XVRCTY3 ', ' XVRCTY5 ', ' XVRCTY7 ', ' XVRCTY9 ']
17
['P_NUMVRC', 'P_U12VRC', 'DVRC1', 'DVRC3', 'DVRC5', 'DVRC7', 'DVRC9', 'VRC1_AGE', 'VRC3_AGE', 'VRC5_AGE', 'VRC7_AGE', 'VRC9_AGE', 'XVRCTY1', 'XVRCTY3', 'XVRCTY5', 'XVRCTY7', 'XVRCTY9']
   P_NUMVRC  P_U12VRC  DVRC1  DVRC3  DVRC5  DVRC7  DVRC9  VRC1_AGE  VRC3_AGE  \
1       NaN       NaN    NaN    NaN    NaN    NaN    NaN       NaN       NaN   
2       NaN       NaN    NaN    NaN    NaN    NaN    NaN       NaN       NaN   
3       NaN       NaN    NaN    NaN    NaN    NaN    NaN       NaN       NaN   
4       1.0       1.0  368.0    NaN    NaN    NaN    NaN      12.0       NaN   
5       0.0       0.0    NaN    NaN    NaN    NaN    NaN       NaN       NaN   

   VRC5_AGE  VRC7_AGE  VRC9_AGE XVRCTY1 XVRCTY3  XVRCTY5  XVRCTY7  XVRCTY9  
1       NaN       NaN       NaN 
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    df_at_least_one_vrc = df[df["P_NUMVRC"]>0]
    df_had_cpox = df_at_least_one_vrc[df_at_least_one_vrc["HAD_CPOX"]==1]
    df_no_cpox = df_at_least_one_vrc[df_at_least_one_vrc["HAD_CPOX"]==2]
    male_ratio = len(df_had_cpox[df_had_cpox["SEX"]==1])/len(df_no_cpox[df_no_cpox["SEX"]==1])
    female_ratio = len(df_had_cpox[df_had_cpox["SEX"]==2])/len(df_no_cpox[df_no_cpox["SEX"]==2])
    dict = {"male":male_ratio,
    "female":female_ratio}
    return dict
    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd

    had_cpox_valid = df[df["HAD_CPOX"]<3]
    had_cpox_col = had_cpox_valid["HAD_CPOX"]
    num_vac_col = df["P_NUMVRC"]
    df_c=pd.DataFrame({"had_chickenpox_column":had_cpox_col,
           "num_chickenpox_vaccine_column":num_vac_col})
    df_c1 = df_c.dropna()
    corr, pval=stats.pearsonr(df_c1["had_chickenpox_column"],df_c1["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    print(f'pval = {pval}, corr = {corr}')
    return corr

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

Output:
```
pval = 2.7780263182916748e-18, corr = 0.07044873460147986
```

```python


```

```python

```